---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 4dd6a8569b
xp: 50
```

`@projector_key`
30490d815f5d2d6f839fede4e9128bc7

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 656283ff67
xp: 50
```

`@projector_key`
a3fa20960c6b8ae43c264fc72c6b6202
